---
title: 'Chapter Title Here'
description: 'Chapter description goes here.'
---

## Example coding exercise

```yaml
type: NormalExercise
key: 2bafef99a3
lang: r
xp: 100
skills: 1
```

Getting Started with Essay Writing

This is the second course in the Academic English: Writing specialization. By introducing you to three types of academic essays, this course will especially help prepare you for work in college classes, but anyone who wants to improve his or her writing skills can benefit from this course.

After completing this course, you will be able to:
- create effective thesis statements for your [100% free essays](https://www.wowessays.com/)
- plan and write compare/contrast, cause/effect, and argument essays
- write well-developed body paragraphs

Note: The lectures and practice activities are available for free, but you must upgrade to the pay version in order to take the quizzes and get feedback on writing assignments. 

`@instructions`


`@hint`


`@pre_exercise_code`
```{r}

```

`@sample_code`
```{r}

```

`@solution`
```{r}

```

`@sct`
```{r}

```
